Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Emil Spåre"
COLLABORATORS = ""
```

---

### Basic start

A Python manual can be found under https://www.python.org; an overview of the possibilities of Python can be found under https://perso.limsi.fr/pointal/python:memento.

#### Exercise 1
Write a function 'max_2()' that accepts two numbers as parameters and that prints the larger number to the screen. For this elementary exercise, do not use the built-in function 'max()' - just revert to elementary expressions and statements.

<!-- grade_id: cell-ee8c56c7cc588b03, points: 0, grade: false, solution: true -->
```python
def max_2(x,y):
    # YOUR CODE HERE
    if x > y:
        return(x)
    else:
        return(y)

a=3
b=5

print("The maximum of %d and %d is %d" % (a, b, max_2(a, b)))
```

Output:
```
The maximum of 3 and 5 is 5
```

<!-- grade_id: cell-304328f9c7d4c051, points: 0, grade: true, solution: false -->
```python
assert max_2(2,6)==6
```

Rewrite the print statement so that the maximum of five numbers is written out: "The maximum of a, b, c, d, and e is...". Use the same 'max_2()' function.   

<!-- grade_id: cell-93a170e9b05f7b24, points: 0, grade: false, solution: true -->
```python
a=3
b=5
c=4
d=6
e=1

# YOUR CODE HERE
def max_2(f, g, h, i, j):
    # YOUR CODE HERE
    x = [f, g, h, i, j]
    x.sort()
    return(x[-1])
            
a=3
b=5

print("The maximum of %d, %d, %d, %d and %d is %d" % (a, b, c, d, e, max_2(a, b, c, d, e)))
```

Output:
```
The maximum of 3, 5, 4, 6 and 1 is 6
```

#### Exercise 2

A function to describe the length of a string 'n' can be given by

```python
def string_length( n ):
    u=len(n)
    return u
```

Write 'len(n)' in terms of elementary expressions and statements.

<!-- grade_id: cell-b607688727ab57f1, points: 0, grade: false, solution: true -->
```python
def string_length2( n ):
    # YOUR CODE HERE
    x = list(n)
    y = 0
    for item in x:
        y += 1
    return(y)

n= 'Obelix did fall into the pot of the druide.'
print ("The length of the string is %d" % string_length2(n))
```

Output:
```
The length of the string is 43
```

<!-- grade_id: cell-24ac435261756312, points: 0, grade: true, solution: false -->
```python
assert string_length2('Exercise is correct')==19
```

#### Exercise 3

A palindrome is a word which is the same word when it is read from left to right as from right to left. Write a function 'is_palindrome()' that checks if a word is a palindrome, e.g. is_palindrome("noon") --> True. 

<!-- grade_id: cell-b555fc0823ca1ecf, points: 0, grade: false, solution: true -->
```python
def is_palindrome( n ):
    # YOUR CODE HERE
    m = list(reversed(n))
    l = "".join(m)
    if l == n:
        return(True)
    else:
        return(False)
is_palindrome("noon")
```

Output:
```
True
```

<!-- grade_id: cell-499ffeac6ffa1eb3, points: 0, grade: true, solution: false -->
```python
assert is_palindrome("rotor")==True
```

A phrase is a palindrome if the word, which is obtained by omitting all non-letters as well as by interchanging all capital letters by their small counterparts, is a palondrome. A famous example of such a phrase is "A man, a plan, a canal: Panama."
Write a program which checks whether a phrase is a palindrome. 

<!-- grade_id: cell-6296442ff4958ebb, points: 0, grade: false, solution: true -->
```python
def is_palindrome_phrase(w):
    # YOUR CODE HERE
    a = w.lower()
    b = list(a)
    c = []
    for item in a:
        if item.isalpha():
            c.append(item)
    d = "".join(c)
    e = "".join(reversed(c))
    if d == e:
        return(True)
    else:
        return(False)

phrase = "A man, a plan, a canal: Panama."    
print ("The phrase '%s' is a palindrome: %s" %  (phrase, is_palindrome_phrase(phrase)))
```

Output:
```
The phrase 'A man, a plan, a canal: Panama.' is a palindrome: True
```

<!-- grade_id: cell-8fb5a5d12a7a6877, points: 0, grade: true, solution: false -->
```python
assert is_palindrome_phrase('Doc, note, I Dissent. A fast never prevents a fatness! I diet on cod.')==True
```

#### Exercise 4

Write a function that tests (True/False) if a number or character 'v' is present in a list of values or a string 'l'.

<!-- grade_id: cell-406e03b5da076705, points: 0, grade: false, solution: true -->
```python
def is_member(v,l):
    # YOUR CODE HERE
    for item in l:
        if item == v:
            return(True)
    return(False)
        

v=5                                                                            
l=[4,3,2,5]

print("Is 5 a member of the list? %s" % (is_member(v,l)))
```

Output:
```
Is 5 a member of the list? True
```

<!-- grade_id: cell-51bd927fdfdc4c8a, points: 0, grade: true, solution: false -->
```python
assert is_member(6,l)==False
```

<!-- grade_id: cell-ca070876e62a4658, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
thing = "x"
container = "panoramix"
print("Is %s a member of %s? %s" % (thing, container, is_member(thing, container)))
```

Output:
```
Is x a member of panoramix? True
```

<!-- grade_id: cell-031caa7cd3eb4e3b, points: 0, grade: true, solution: false -->
```python
assert is_member("x","panoramix")==True
```

#### Exercise 5

Write a function 'generate_string(n,c)' which generates a string which consists out of 'n' times the characther 'c'. Use a loop and add each time a new character 'c' to the string.

<!-- grade_id: cell-bdafd275123a0ab4, points: 0, grade: false, solution: true -->
```python
def generate_string(n,c):
    # YOUR CODE HERE
    thing = ""
    while n > 0:
        thing += c
        n -= 1
    return(thing)

generate_string(6,'k')
```

Output:
```
'kkkkkk'
```

<!-- grade_id: cell-481d83412f434b99, points: 0, grade: true, solution: false -->
```python
assert generate_string(3,'s')=='sss'
```

Write a function which prints a string of 'y' for each entry of a vector [2,4,1]:

yy      
yyyy  
y


<!-- grade_id: cell-631ae09a389a66d4, points: 0, grade: false, solution: true -->
```python
def star_histogram(c):
    # YOUR CODE HERE
    y = []
    for item in c:
        x = ""
        while item > 0:
            x += "y"
            item -= 1
        y.append(x)
    y = "\n".join(y) + "\n"
    return(y)
    
c=[2,4,1]
print(star_histogram(c))
```

Output:
```
yy
yyyy
y
```

<!-- grade_id: cell-901ea19debe3ee4c, points: 0, grade: true, solution: false -->
```python
assert star_histogram(c)=="yy\nyyyy\ny\n"
```

#### Exercise 6

Compute how many times a character is present in a string. Use a dictionary to represent the string. 

<!-- grade_id: cell-94b1ecc182583f9b, points: 0, grade: false, solution: true -->
```python
def character_frequency(word):
    # YOUR CODE HERE
    thing = {}
    for item in word:
        thing[item] = word.count(item)
    return(thing)
        
c="abababbca"
print("How many times a character is present: %s" % character_frequency(c))
```

Output:
```
How many times a character is present: {'a': 4, 'b': 4, 'c': 1}
```

<!-- grade_id: cell-2a787760db753909, points: 0, grade: true, solution: false -->
```python
assert character_frequency("abra")=={'a': 2, 'b': 1, 'r': 1}
```

### Exercise 7

Make a list 'mult_7' of all the integer numbers from 0 to 99 which are multiples of 7. Print these values to the screen.

<!-- grade_id: cell-8091f8ddc1391ba5, points: 0, grade: false, solution: true -->
```python
def mult_7():
    # YOUR CODE HERE
    thing = []
    for item in range(99):
        if item % 7 == 0:
            thing.append(item)
    return(thing)
print("The list is %s" % mult_7() )
```

Output:
```
The list is [0, 7, 14, 21, 28, 35, 42, 49, 56, 63, 70, 77, 84, 91, 98]
```

<!-- grade_id: cell-825e1592f0e1efbb, points: 0, grade: true, solution: false -->
```python
assert mult_7()==[0, 7, 14, 21, 28, 35, 42, 49, 56, 63, 70, 77, 84, 91, 98]
```

```python

```