Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Isac Ingfeldt"
COLLABORATORS = "None"
```

---

# Using decorators in Python

### A. Function variables: 

Create a generic function integrate_func(func, n), which implements arbitrary function integration in interval [-1,1] 
using Chebyshev–Gauss quadrature of first kind: 
 
$\int_{-1}^{+1} {\frac{f(x)}{\sqrt{(1-x^2)}}dx} \approx \sum_{i=1}^{n} \omega_i f(x_i)$ 

where the i-th grid point weight and values are    

$x_i = cos(\frac{2i-1}{2n} \pi)$ and  $\omega_i = \frac{\pi}{n}$. 

* Function integrate_func(func, n):
* Parameters
    - func is the mathematical function, which will be integrated. 
    - n is the number of grid points. 
* Return
    - the value of integral (single float number).

<!-- grade_id: cell-106104681ce80401, points: 0, grade: false, solution: true -->
```python
import math
def integrate_func(func, n):
    omg = (math.pi / n)
    sum = 0
    for i in range(n):
        x = math.cos(((2 * i) - 1) / (2 * n) * math.pi)
        f = func(x) * omg
        sum += f
    print('Integral :', sum)
    return sum
```

<!-- grade_id: cell-e8becee65f92a575, points: 1, grade: true, solution: false -->
```python
# Verify integration function 

def one_func(x): 
    return math.sqrt(1.0 - x * x)

assert math.fabs(integrate_func(one_func, 100) - 2.0) < 1.0e-3

assert math.fabs(integrate_func(math.cos, 100) - 2.403939430634413) < 1.0e-4
```

Output:
```
Integral : 2.0000822490709877
Integral : 2.403939430634413
```

### B. Simple decorators 

Create decorators for an arbitrary mathematical function f(x):
* decorator square, which applies $1 - f(x)^2$ transformation to f(x).
* decorator gaussian, which applies $exp(-0.5 * f(x)^2)$ transformation to f(x).

Create decorated functions (using @decorator syntax): 
* cos_func(x), which computes cosine of x, and decorate it with square decorator.
* gau_func(x), which computes 2.0 * x, and decorate it with gaussian decorator. 
* comp_func(x), which computes 2.0 * x, and decorate it with square and gaussian decorators,  such that a function $f(x)$ is modified to 
$$1-\left(exp(-0.5*f(x)^2)\right)^2$$   

<!-- grade_id: cell-dbe1d3027a05e5a4, points: 0, grade: false, solution: true -->
```python
import math

#------------------------Decorators
def square(func):
    def squaring(x):
        return(1 - (func(x) * func(x)))
    return(squaring)

def gaussian(func):
    def gauss(x):
        return(math.exp(-0.5 * (func(x)*func(x))))
    return(gauss)


#-----------------------Math functions
@square
def cos_func(x):
    ans = math.cos(x)
    return(ans)

@gaussian
def gau_func(x):
    ans = 2.0 * x
    return(ans)

@square
@gaussian
def comp_func(x):
    ans = 2.0 * x
    return(ans)
```

<!-- grade_id: cell-73e7fae89242df32, points: 1, grade: true, solution: false -->
```python
#Verify square decorator for cos_func

assert cos_func(0.0) == 0.0 

assert cos_func(math.pi) == 0.0

assert cos_func(math.pi*0.5) == 1.0 

assert math.fabs(cos_func(math.pi*0.25) - 0.5) < 1.0e-6

```

<!-- grade_id: cell-4181fe5f998c9a96, points: 1, grade: true, solution: false -->
```python
#Verify gaussian decorator for gau_func

assert gau_func(0.0) == 1.0 

assert math.fabs(gau_func(1.0) - 0.135335283236613) < 1.0e-6 

assert math.fabs(gau_func(0.5) - 0.606530659712633) < 1.0e-6 
```

<!-- grade_id: cell-fa09de3f68f4fad3, points: 1, grade: true, solution: false -->
```python
#Verify comp_func

assert comp_func(0.0) == 0.0 

comp_func(0.5)

assert math.fabs(comp_func(0.5) - 0.632120558828558) < 1.0e-6 

assert math.fabs(comp_func(0.3) - 0.302323673928969) < 1.0e-6 
```

### C. A timer decorator

In this exercise we write a decorator to report the time it takes to complete a function. Hardcoding a timer like this could look like

```python
import time
def function_taking_time():
    t1 = time.time()
    time.sleep(3)
    t2 = time.time()
    print("Time spent in function_taking_time: {:.1f} seconds.".format(t2-t1))  
```

```python
function_taking_time()
```

Output:
```
Time spent in function_taking_time: 3.0 seconds.
```

Now write a decorator `timeme` which accomplishes this for any function taking any arguments. Instead of printing to the screen write the timing info to a logfile with the same name as the function and a `.log` extension.

Hint: make use of the `__name__` attribute of a function.

<!-- grade_id: cell-b478018e06675d4e, points: 0, grade: false, solution: true -->
```python
import time
def timeme(method):
    "Timer decorator"
    def timed(*args, **kwargs):
        tstart = time.time()
        result = method(*args, **kwargs)
        tend = time.time()
        telapse = (tend - tstart)
        telapse = str(format(telapse, '.1f'))
        
        name = method.__name__
        log = open(name + '.log', 'w')
        string = ('Time spent in slow_append: ' + telapse + ' seconds' )
        log.write(string)
        print(string)
        return result
    return timed
```

The function below appends `x` to a given list `l` if it is given as an input, otherwise it does nothing. The reported time spend in the function will reflect this when applying the completed decorator

```python
@timeme
def slow_append(l, x=None):
    if x is not None:
        time.sleep(3)
        l.append(x)
```

<!-- grade_id: cell-4512178e73c67be5, points: 1, grade: true, solution: false -->
```python
l = []
slow_append(l)
assert open('slow_append.log').read() == "Time spent in slow_append: 0.0 seconds"
assert l == []
```

Output:
```
Time spent in slow_append: 0.0 seconds
```

<!-- grade_id: cell-e3b9db4ad3bedd1c, points: 1, grade: true, solution: false -->
```python
l = [1, 2]
slow_append(l, x=3)
assert open('slow_append.log').read() == "Time spent in slow_append: 3.0 seconds"
assert l == [1, 2, 3]
```

Output:
```
Time spent in slow_append: 3.0 seconds
```

```python

```

```python

```