Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Isac Ingfeldt"
COLLABORATORS = "None"
```

---

### Linux & bash

In this exercise, we focus upon the most important linux and bash commands. However, a whole new universe is out there - the reader is advised to discover more about this virtual environment on his/her own initiative.  

In order to facilitate the exercises of today, let us first prepare some files, which will be used later on to train the concepts and commands discussed in the lecture. Simply execute the next block by using shift+enter.

```python
%%bash
echo '2 wild boars' > breakfast_obelix
echo '1 duck' >> breakfast_obelix
echo '2 kg potatoes' >> breakfast_obelix
echo 'print("I am the emperor of Rome")' > caesar
```

Let us now start the exercise. 

In each block in Jupyter, the initial line %%bash has to be written: it tells the operating system that the lines to follow will be interpreted as bash commands. Since all blocks are technically independent from each other, it is important to write this command each time explicitly.

Check where you are in the file system at this moment. List all files. 

<!-- grade_id: cell-db71c50869567568, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
pdw
ls
```

Make an architecture based upon 3 folders ('asterix', 'obelix', 'romans'). Inside the folder 'asterix', create an empty file 'breakfast'. Inside the folder 'obelix', create a directory 'latraviata'. Move the file 'breakfast_obelix' to the folder 'obelix'.

<!-- grade_id: cell-f15ded1957c3b1f4, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
mkdir asterix
mkdir obelix
mkdir romans

cd asterix
touch breakfast
cd ..

cd obelix
touch latraviata

mv breakfast_obelix obelix/breakfast_obelix
```

<!-- grade_id: cell-e0d969346040a77d, points: 0.0, grade: true, solution: false -->
```python
import os.path
assert os.path.isdir('asterix')
assert os.path.isdir('obelix')
assert os.path.isdir('romans')
assert os.path.isdir('obelix/latraviata')
assert os.path.isfile('obelix/breakfast_obelix')
assert os.path.isfile('asterix/breakfast')
```

Inside the folder 'obelix', move 'breakfast_obelix' to 'latraviata' and change the name of the file 'breakfast_obelix' into 'middag_with_family'. Move now the folder 'latraviata' from 'obelix' to 'asterix'.

<!-- grade_id: cell-0d9b30453d6c1694, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
cd obelix

```

<!-- grade_id: cell-a5819be950b09a04, points: 0.0, grade: true, solution: false -->
```python
import os.path
assert os.path.isdir('asterix/latraviata')
assert os.path.isfile('asterix/latraviata/middag_with_family')
```

Make a tree to verify where you are now in the architecture of the files. Remove the directory 'romans' and remove the file breakfast in 'asterix'.

<!-- grade_id: cell-1672cfbb11914e59, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
```

<!-- grade_id: cell-ff17c5fa65d00d5d, points: 0.0, grade: true, solution: false -->
```python
import os.path
para=False
if not os.path.isfile('asterix/breakfast'):
    para=True
assert para==True
```

Copy the file 'asterix/latraviata/middag_with_family' to 'obelix' and call it again 'breakfast'. Use 'grep' to find out how many wild boars were included in 'breakfast'. 

<!-- grade_id: cell-382595e252ed681d, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
```

<!-- grade_id: cell-2d90598e190d6a49, points: 0.0, grade: true, solution: false -->
```python
import os.path
assert os.path.isfile('obelix/breakfast')
```

Inside the folder 'obelix', create a folder 'idefix' and inside 'idefix' create the empty files 'pup1', 'pup2' and 'pup3'. Go to the folder of obelix and show the time in which the content of this folder (being the folder 'idefix' and the file 'breakfast') has been made. Navigate finally one folder higher and use the keyword 'find' to search where 'pup1' is located.

<!-- grade_id: cell-e2929ef5e2e1aa0e, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
```

<!-- grade_id: cell-4215aa1773e6106f, points: 0.0, grade: true, solution: false -->
```python
import os.path
assert os.path.isdir('obelix/idefix')
assert os.path.isfile('obelix/idefix/pup1')
assert os.path.isfile('obelix/idefix/pup2')
assert os.path.isfile('obelix/idefix/pup3')
```

The file 'caesar' is in fact a text file, which contains python code. Show the content of the file. Run the program using python3.

<!-- grade_id: cell-40da915435f48a13, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
```

```python

```